# Ex01 - Working with Jupyter notebook
This goal exercise is in order to practice the use of Jupyter notebook and understand the way the coding assignments are tested.

## Working on the assignment, todo and not todo:
<b><u>Todo</u></b>:<br />
<div class="alert alert-success">
<ul><li>✔ <b> Work by yourself</b> and submit your own assignment, <b>no pairing</b> to other students</li>
<li>✔ Test and save your assignment - <b>submit the last tested and saved version</b></li>
<li>✔ Work on the <b>original notebook</b> of the assignment</li>
<li>✔ submit <b>only</b> the <b>notebook (.ipynb)</b> file</li></ul>
</div>
<b><u>Not todo</u></b>:<br />
<div class="alert alert-danger">
<ul><li>✘ do <b>NOT</b> submit an <b>empty assignment</b></li>
<li>✘ do <b>NOT submit extra files</b>, <u>unless you're asked to do so</u></li>
<li>✘ <b>Do NOT submit a .py/.txt/.rar/.zip (or any non (.ipynb) file)</b> <u>version for the notebook</u> of the assignment</li>
<li>✘ <b>do NOT change</b> the notebook file name</li></ul>
</div>

# Answer code cells:
<div class="alert alert-info">
<ul><li>Put your code one line after the '<code># YOUR CODE HERE</code>' remark</li>
<li>You must remove the '<code>raise NotImplementedError()</code>' exception raising line (which appears a line after the above remark).</li>
<ul><li>if you do not remove this line, it will be a sign that you didn't implement that code, and we won't be able to check your work</li></ul>
<li>Do NOT remove any other line in the notebook</li></ul>
</div>

### An example code cell:
<div class="alert alert-info"><code># --------------------------- (AFTER YOUR IMPLEMENTATION) RUN THIS CODE CELL  -------------------------------------
'''
What do you need to do?

Complete the 'method_to_implement' function to do something (some explanation ...) and
     return that value back
'''
def method_to_implement():
    # YOUR CODE HERE
    raise NotImplementedError()
</code></div>

<h3><u>Not to do</u>:</h3><br />
<div class="alert alert-danger">
The following is <b>not good</b>, the '<code>raise NotImplementedError()</code>' line was <u>not removed</u>. <br />
Notes:<br />
<ul><li>Such a case will <b>cause an exception</b>, even if you implemented the function</li>
<li>You will <b>not get any points</b> for your implementation, even if you implemented the function</li></ul>
</div>

<b>Case 1 - Not good</b>:<br/>
<div class="alert alert-danger">
<ul><li>✘ Code not Implemented</li>
<li>✘ No returned value</li>
<li>✘ Did not remove the '<code>raise NotImplementedError()</code>' line</li></ul>
<code># --------------------------- (AFTER YOUR IMPLEMENTATION) RUN THIS CODE CELL  -------------------------------------
'''
What do you need to do?

Complete the 'method_to_implement' function to do something (some explanation ...) and
     return that value back
'''
def method_to_implement():
    # YOUR CODE HERE
    raise NotImplementedError()
</code></div>

<b>Case 2 - Still not good</b>:<br/>
<div class="alert alert-danger">
<ul><li>✔ Implemented code</li>
<li>✔ returned expected value</li>
<li>✘ Did not remove the '<code>raise NotImplementedError()</code>' line</li></ul>
<code># --------------------------- (AFTER YOUR IMPLEMENTATION) RUN THIS CODE CELL  -------------------------------------
'''
What do you need to do?

Complete the 'method_to_implement' function to do something (some explanation ...) and
     return that value back
'''
def method_to_implement():
    # YOUR CODE HERE
    raise NotImplementedError()
    implemented code ...
    return val
</code></div>

<b>Case 3 - Still not good</b>:<br/>
<div class="alert alert-danger">
<ul><li>✔ Removed the '<code>raise NotImplementedError()</code>' line</li>
<li>✔ Implemented code</li>
<li>✘ No returned value, <u>though expected</u></li></ul>
<code># --------------------------- (AFTER YOUR IMPLEMENTATION) RUN THIS CODE CELL  -------------------------------------
'''
What do you need to do?

Complete the 'method_to_implement' function to do something (some explanation ...) and
     return that value back
'''
def method_to_implement():
    # YOUR CODE HERE
    implemented code ...
</code></div>

<b>Case 4 - Good</b>:<br/>
<div class="alert alert-success">
<ul><li>✔ Removed the '<code>raise NotImplementedError()</code>' line</li>
<li>✔ Implemented code</li>
<li>✔ returned expected value</li></ul>
<code># --------------------------- (AFTER YOUR IMPLEMENTATION) RUN THIS CODE CELL  -------------------------------------
'''
What do you need to do?

Complete the 'method_to_implement' function to do something (some explanation ...) and
     return that value back
'''
def method_to_implement():
    # YOUR CODE HERE
    implemented code ...
</code></div>

## Preceding Step - import packages
This step is necessary in order to use external packages. <br/><br/>

```python
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# === CELL TYPE: IMPORTS AND SETUP 

import os                     # for testing use only
from datetime import datetime # for testing use only
import pandas as pd
import numpy as np
```

<h3>Part 1 - Student information methods:</h3><br/>
The following 2 cells consist of 2 methods which aim to validate your details.<br/>
Please <b>make sure to implement <u>both</u> of them</b><br/><br/>
<table align='left'>
    <tr><td>
        <img src="./images/id-card.png" alt="name and id" width="150" align='left'/>
    </td></tr>
</table>    

<u>Part 1.a. Student information methods</u>:
<u>method name</u>: <b>myName</b>
<pre>The following is expected:
--- the method needs to return your full name (as a string, in English)
    For example: assume your name is John Smith, you should write:
    return 'John Smith'
------------
return value:
- your full name (as a string), in English
---------------------
</pre>

<!-- grade_id: answer_1a_myName, points: 0, grade: false, solution: true -->
```python
# 1.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ANSWER 

def myName():
    # ADD the return statement in the LINE AFTER: '# YOUR CODE HERE'
    # YOUR CODE HERE
    return 'Haim'
```

```python
# 1.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: SELF TESTING
# Use the following code to test your implementation:

student_name = myName()
# YOUR CODE HERE

```

Output:
```
Haim
```

```python

```

```python
# 1.a.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation, if used)
# === CODE TYPE: SELF TESTING
# Add your additional tests here if needed:

# YOUR CODE HERE

```

<u>Part 1.b. Student information methods</u>:<br/>
<u>method name</u>: <b>myId</b><br/>
<pre>
The following is expected:
--- the method needs to return your ID number (as an integer number)
    For example: assume your ID number is 1234, you should write:
    return 1234
    * Note: the ID must fit your ID in Moodle. 
------------
return value:
- your ID number (as an integer number)
------------
</pre>

<!-- grade_id: answer_1b_myId, points: 0, grade: false, solution: true -->
```python
# 1.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: ANSWER 

def myId():
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# 1.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation)
# === CODE TYPE: SELF TESTING
# Use the following code to test your implementation:

student_id = myId()
# YOUR CODE HERE

```

```python
# 1.b.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run after implementation, if used)
# === CODE TYPE: SELF TESTING
# Add your additional tests here if needed:

# YOUR CODE HERE

```

<!-- grade_id: test1_1-student-info, points: 0.5, grade: true, solution: false -->
```python
# 1.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 1 (1.a., 1.b.) - Test 1 (... points)")
print ("\t---> Testing the implementation of 'myName' and 'myId' ...\n")

# dataframe for output:
dt1 = datetime.now()
df_output = pd.DataFrame(index=['student_name','student_id'],columns=['value', 'type', 'date_time'])
out_file_name = 'a0_jupyter_notebook_assignment_output.csv'

try:
    aName = myName()
    aId = myId()
except Exception as e:
    print ('You probably have a syntax or implementation error, we got the following exception:')
    print ('\tError Message:',str(e))
    print ('Try fixing your implementation')
    raise     
    
assert aName is not None, 'no return student name'
assert aId is not None, 'no return student id'

df_output['value'] = [aName, aId]
df_output['type']  = [type(aName), type(aId)]
dt2 = datetime.now()
df_output['date_time']  = [dt1.strftime('%Y-%m-%d %H:%M:%S'), dt2.strftime('%Y-%m-%d %H:%M:%S')]
df_output.to_csv(out_file_name)
print ('Output dataframe:')
print ('-----------------')
print (df_output)

print ("\nGood Job!\nYou've passed the 1st test for the implementation of 'myName' and 'myId'  :-)")

```

<!-- grade_id: test1_2-student-info, points: 0.5, grade: true, solution: false -->
```python
# 1.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: GRADED TEST 

print ("Part 1 (1.a., 1.b.) - Test 2 (... points)")
print ("\t---> Testing the implementation of 'myName' and 'myId' ...\n")

try:
    aName = myName()
    aId = myId()
except Exception as e:
    print ('You probably have a syntax or implementation error, we got the following exception:')
    print ('\tError Message:',str(e))
    print ('Try fixing your implementation')
    raise 
    
assert type(aName) is str, "name is not a string: %r" % aName
assert type(aId) is int, "id is not an integer: %r" % aId

print ("\nGood Job!\nYou've passed the 2nd test for the implementation of 'myName' and 'myId'  :-)")
```

```python
# 1.
# ------------>>>>>>>> RUN THIS CODE CELL <<<<<<<<------------
# --------  (run only)
# === CODE TYPE: INFO

print ("Part 1 (1.a., 1.b.) - Test 2 (1 point)")
print ("\t---> Testing the implementation of 'myName' and 'myId' ...\n")

print ("\nGood Job!\nAdditional testing will be done  :-)")
```